```python
# Initialize Otter
import otter
grader = otter.Notebook("test1.ipynb")
```

<a href="https://colab.research.google.com/github/hernansalinas/autogradesMetodosComputacionales/blob/main/Actividades_clase/actividad03/test1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>



```python
import otter
import pandas as pd
import numpy as np
%matplotlib inline
import otter
grader = otter.Notebook()
```

**Question 1:** Escribir una función llamada `promedio` que retorne el promedio de una arrray de numpy 

```python
def promedio(x):
    ...
```

```python
grader.check("q1")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: None)
```
>>> x_ = np.arange(1, 100)
>>> promedio(x_) == x_.mean()
True
```
